# Домашнее задание №3. Вариант рекомендательной системы с knn

```python
import os
import os.path
import warnings
from collections import Counter


from typing import Union


from datetime import timedelta

import implicit
import rectools


import requests


import scipy as sp


import numpy as np


import pandas as pd


from rectools import Columns


from rectools.metrics import MAP, MeanInvUserFreq


from rectools.dataset import Dataset


from scipy.sparse import coo_matrix, spmatrix


from implicit.nearest_neighbours import ItemItemRecommender



warnings.filterwarnings("ignore")
```

```python
data_path = os.environ.get("DATA_PATH")
```

```python
print("implicit:", implicit.__version__)
print("requests:", requests.__version__)
print("rectools:", rectools.__version__)
print("pandas:", pd.__version__)
print("numpy:", np.__version__)
print("scipy:", sp.__version__)
```

Output:
```
implicit: 0.7.2
requests: 2.32.3
rectools: 0.12.0
pandas: 2.2.3
numpy: 1.26.4
scipy: 1.12.0
```

```python
%%time

!wget -q https://github.com/irsafilo/KION_DATASET/raw/f69775be31fa5779907cf0a92ddedb70037fb5ae/data_original.zip -O kion_train.zip
!unzip -o kion_train.zip -x '__MACOSX/*'
!rm kion_train.zip
```

```python
if data_path is None:
    data_path = "./data_original"
```

```python
interactions = pd.read_csv(os.path.join(data_path, "interactions.csv"), parse_dates=["last_watch_dt"])
interactions = interactions.rename(columns={"total_dur": Columns.Weight, "last_watch_dt": Columns.Datetime})


users = pd.read_csv(os.path.join(data_path, "users.csv"))
items = pd.read_csv(os.path.join(data_path, "items.csv"))


print(interactions.shape)
interactions.head(5)
```

Output:
```
(5476251, 5)
   user_id  item_id   datetime  weight  watched_pct
0   176549     9506 2021-05-11    4250         72.0
1   699317     1659 2021-05-29    8317        100.0
2   656683     7107 2021-05-09      10          0.0
3   864613     7638 2021-07-05   14483        100.0
4   964868     9506 2021-04-30    6725        100.0
```

# Ситуация: 

Ваш коллега ушел в отпуск, оставив вам недописанную модель в `userknn.py` (код из семинара №3 будет в ячейке этого ноутбука далее)

Через 3 дня ее нужно выкатить в А/B тест. Это значит, что модель должна уметь выдавать рекомендации размера `k` любым пользователям (холодным / горячим). 

Требуется дописать методы класса UserKnn:

- метод, который выдает рекомендации холодным пользователям (сейчас модель этого не умеет)
- в методе, который выдает рекомендации горячим пользователям, дописать код так, чтобы у всех было одинаковое кол-во рекомендаций k, а не меньше k, как сейчас.

Также была договоренность, что перед выкаткой вы покажете продакт менеджеру финальные значения по оффлайн метрикам.

# Готовим тестовую выборку

Здесь и далее жестко фиксируем тестовую выборку **ровно последние 14 дней из `interactions`**  - по результатам на ней будут выставляться баллы

Сохраните тестовую выборку (тестовые interactions) в pandas DataFrame и назовите его `test`. 

Оставьте в нем все поля из `interactions`. Удостоверьтесь, что формат полей следующий:

    Column   Dtype         
 - user_id -    int64         
 - item_id -    int64         
 - datetime -   datetime64
 - weight -      int64         
 - watched_pct - float64    


<!-- grade_id: cell-d6819c2c430a8e75, points: 0, grade: false, solution: true -->
```python
max_date = interactions[Columns.Datetime].max()
test = interactions[interactions[Columns.Datetime] >= (max_date - pd.Timedelta(days=13))].copy()
train = interactions[interactions[Columns.Datetime] < (max_date - pd.Timedelta(days=13))].copy()
print(test.shape)
print(train.shape)
```

Output:
```
(827089, 5)
(4649162, 5)
```

```python
for df in [test, train]:
    df["user_id"] = df["user_id"].astype("int64")
    df["item_id"] = df["item_id"].astype("int64")
    df[Columns.Datetime] = pd.to_datetime(df[Columns.Datetime])
    df[Columns.Weight] = df[Columns.Weight].astype("int64")
    df["watched_pct"] = df["watched_pct"].astype("float64")
```

```python
# test.to_csv(data_base_path / "test.csv", index=False)
# train.to_csv(data_base_path / "train.csv", index=False)
```

## Проверка:  тестовая выборка - 3 балла

Внимание! Есть скрытые тесты 

<!-- grade_id: cell-d4177ea915281721, points: 3, grade: true, solution: false -->
```python
assert isinstance(test, pd.DataFrame)
assert test.datetime.nunique() == 14

expected_mean = 6887.2
actual_mean = round(test["weight"].mean(), 1)
assert abs(actual_mean - expected_mean) < 0.001
```

# Код модели

Допишите методы `fit_cold_model` и  `recommend_cold`, которые делают рекомендации холодным пользователям (сейчас модель этого не умеет)

<!-- grade_id: cell-fc1beb2c90026cfd, points: 0, grade: false, solution: true -->
```python
class UserKnn:
    """
    A user-based KNN model wrapper around `implicit.nearest_neighbours.ItemItemRecommender`
    """

    SIMILAR_USER_COLUMN = "similar_user_id"
    SIMILARITY_COLUMN = "similarity"
    IDF_COLUMN = "idf"

    def __init__(self, model: ItemItemRecommender, N_similar_users: int):
        self.model = model
        self.N_similar_users = N_similar_users

        self.users_inv_mapping = None
        self.users_mapping = None
        self.items_inv_mapping = None
        self.items_mapping = None

        self.watched_items_dataframe = None
        self.item_idf = None
        self.cold_model_fitted = False
        self.cold_items = None

    def _set_mappings(self, interactions: pd.DataFrame) -> None:
        """
        Create dictionaries to map external IDs (users, items) to internal IDs and vice versa.
        """
        unique_users = interactions[Columns.User].unique()
        self.users_inv_mapping = dict(enumerate(unique_users))
        self.users_mapping = {v: k for k, v in self.users_inv_mapping.items()}

        unique_items = interactions[Columns.Item].unique()
        self.items_inv_mapping = dict(enumerate(unique_items))
        self.items_mapping = {v: k for k, v in self.items_inv_mapping.items()}

    def _get_user_item_matrix(self, interactions: pd.DataFrame) -> spmatrix:
        """
        Construct a sparse user-item matrix in CSR format.
        Rows represent users, and columns represent items.
        """
        user_idx = interactions[Columns.User].map(self.users_mapping.get)
        item_idx = interactions[Columns.Item].map(self.items_mapping.get)
        data = interactions[Columns.Weight].astype(np.float64)

        user_item_coo = coo_matrix((data, (user_idx, item_idx)))
        return user_item_coo.tocsr()

    def _set_interacted_items_dataframe(self, interactions: pd.DataFrame) -> None:
        """
        Groups interactions by user to get item_id list for each user
        """
        self.interacted_items_dataframe = (
            interactions.groupby(Columns.User, as_index=False)
            .agg({Columns.Item: list})
            .rename(columns={Columns.User: self.SIMILAR_USER_COLUMN})
        )

    @staticmethod
    def idf(n: int, x: float):
        """
        Calculates IDF for one item
        """
        return np.log((1 + n) / (1 + x) + 1)

    def _count_item_idf(self, interactions: pd.DataFrame) -> None:
        """
        Calculate IDF values for all items present in the interactions dataset
         and store the result in self.item_idf.
        """
        item_freqs = Counter(interactions[Columns.Item].values)
        item_idf_df = pd.DataFrame.from_dict(item_freqs, orient="index", columns=["doc_freq"]).reset_index()
        total_interactions = len(interactions)
        item_idf_df[self.IDF_COLUMN] = item_idf_df["doc_freq"].apply(lambda x: self.idf(total_interactions, x))
        self.item_idf = item_idf_df

    def _prepare_for_model(self, train_interactions: pd.DataFrame) -> None:
        """
        Sets mappings, grouped interactions, calculates idf
        """
        self._set_mappings(train_interactions)
        self._set_interacted_items_dataframe(train_interactions)
        self._count_item_idf(train_interactions)

    def fit_cold_model(self, train_interactions: pd.DataFrame) -> None:
        popularity = train_interactions.groupby(Columns.Item)[Columns.Weight].sum().reset_index()
        popularity = popularity.sort_values(Columns.Weight, ascending=False)
        popularity = popularity.rename(columns={Columns.Weight: Columns.Score})
        self.cold_items = popularity
        self.cold_model_fitted = True

    def recommend_cold(self, users: list, k: int = 100) -> pd.DataFrame:
        if not self.cold_model_fitted:
            raise ValueError("Cold model not fitted. Call fit_cold_model first.")

        top_k = self.cold_items.head(k)

        recs_list = []
        for user in users:
            for rank, (_, row) in enumerate(top_k.iterrows(), start=1):
                recs_list.append(
                    {
                        Columns.User: user,
                        Columns.Item: row[Columns.Item],
                        Columns.Score: row[Columns.Score],
                        Columns.Rank: rank,
                    }
                )

        return pd.DataFrame(recs_list)

    def fit(self, train_interactions: pd.DataFrame) -> None:
        """
        Fit the model on the provided training data.

        Internally:
        1) Prepare mappings, watchlist DataFrame, and item IDF.
        2) Create a user-item matrix and fit the underlying Implicit model.
        """
        self._prepare_for_model(train_interactions)
        user_item_matrix = self._get_user_item_matrix(train_interactions)
        self.model.fit(user_item_matrix.T)
        self.fit_cold_model(train_interactions)

    def _get_similar_users(self, external_user_id: int) -> tuple[list[int], list[float]]:
        """
        Retrieve a list of similar users and corresponding similarities
        from the underlying Implicit model.
        """
        if external_user_id not in self.users_mapping:
            # if user doesn't exist in mapping, return sentinel (-1).
            return [-1], [-1]

        internal_user_id = self.users_mapping[external_user_id]
        user_ids, similarities = self.model.similar_items(internal_user_id, N=self.N_similar_users)
        # convert back to external IDs
        external_user_ids = [self.users_inv_mapping[u_id] for u_id in user_ids]
        return external_user_ids, similarities

    @staticmethod
    def get_rank(recs: pd.DataFrame, k: int) -> pd.DataFrame:
        """
        Sort recommendations by score in descending order,
        assign ranks within each user group, and then truncate by top-k.
        """
        recs = recs.sort_values([Columns.User, Columns.Score], ascending=False)
        recs = recs.drop_duplicates([Columns.User, Columns.Item])
        recs[Columns.Rank] = recs.groupby(Columns.User).cumcount() + 1
        recs = recs[recs[Columns.Rank] <= k][[Columns.User, Columns.Item, Columns.Score, Columns.Rank]]

        return recs

    def recommend(self, users: np.ndarray, k: int) -> pd.DataFrame:
        """
        Generate top-k recommendations for the specified list of users.

        Steps:
        1) Find similar users for each target user.
        2) Join watched items from these similar users.
        3) Compute a final score as similarity * IDF.
        4) Return top-k items per user.
        """

        recs = pd.DataFrame({Columns.User: users})
        recs[self.SIMILAR_USER_COLUMN], recs[self.SIMILARITY_COLUMN] = zip(
            *recs[Columns.User].map(lambda user_id: self._get_similar_users(user_id))
        )
        recs = recs.set_index(Columns.User).apply(pd.Series.explode).reset_index()

        recs = (
            recs[~(recs[Columns.User] == recs[self.SIMILAR_USER_COLUMN])]
            .merge(
                self.interacted_items_dataframe,
                on=[self.SIMILAR_USER_COLUMN],
                how="left",
            )
            .explode(Columns.Item)
            .sort_values([Columns.User, self.SIMILARITY_COLUMN], ascending=False)
            .drop_duplicates([Columns.User, Columns.Item], keep="first")
            .merge(self.item_idf, left_on=Columns.Item, right_on="index", how="left")
        )

        recs[Columns.Score] = recs[self.SIMILARITY_COLUMN] * recs[self.IDF_COLUMN]
        recs = recs[[Columns.User, Columns.Item, Columns.Score]]

        existing_counts = (
            recs.groupby(Columns.User)[Columns.Item]
            .nunique()
            .reset_index()
            .rename(columns={Columns.Item: "existing_count"})
        )

        users_df = pd.DataFrame({Columns.User: users})
        recs_count = users_df.merge(existing_counts, on=Columns.User, how="left")
        recs_count["existing_count"] = recs_count["existing_count"].fillna(0).astype(int)

        cold_recs_all = []
        for _, row in recs_count.iterrows():
            user = row[Columns.User]
            count = row["existing_count"]
            if count < k:
                cold_recs = self.recommend_cold([user], k=k)
                existing_items = recs.loc[recs[Columns.User] == user, Columns.Item].unique()
                cold_recs = cold_recs[~cold_recs[Columns.Item].isin(existing_items)]
                cold_recs = cold_recs.head(k - count)
                cold_recs_all.append(cold_recs)

        if cold_recs_all:
            cold_recs_df = pd.concat(cold_recs_all, ignore_index=True)
            recs = pd.concat([recs, cold_recs_df], ignore_index=True)

        recs = self.get_rank(recs, k=k)
        return recs
```

# Рекомендации для холодных пользователей

Используйте метод класса UseKnn, который выдает рекомендации холодным пользователям. 

### Порекомендуйте k=10 уникальных айтемов для каждого холодного пользователя из тестовой выборки. Сохраните рекомендации в pandas DataFrame и назовите его `reco_cold`. 

Датафрейм `reco_cold` должен иметь обязательные поля `user_id`, `item_id`, `rank`. 

⚠️ Холодными считаем пользователей, у которых нет интеракций в train. 


<!-- grade_id: cell-17c883b2dac9bb3d, points: 0, grade: false, solution: true -->
```python
train_users = set(train["user_id"].unique())
test_users = set(test["user_id"].unique())
cold_users = list(test_users.difference(train_users))

print("Number of cold users:", len(cold_users))
```

Output:
```
Number of cold users: 111690
```

```python
user_knn = UserKnn(model=ItemItemRecommender(), N_similar_users=50)
```

```python
user_knn.fit(train)
```

Output:
```
100%|██████████| 850489/850489 [05:49<00:00, 2434.13it/s]
```

```python
# import pickle

# models_base_path = Path("../models/").resolve()

# with open(models_base_path / "user_knn_model.pkl", "wb") as f:
#     pickle.dump(user_knn, f)

# with open(models_base_path / "user_knn_model.pkl", "rb") as f:
#     user_knn = pickle.load(f)
```

```python
reco_cold = user_knn.recommend_cold(cold_users, k=10)
```

```python
# reco_cold.to_csv(data_base_path / "reco_cold.csv", index=False)
```

## Проверка рекомендаций по холодным пользователям - 4 балла

Внимание! Есть скрытые тесты 

<!-- grade_id: cell-e5bf7a9e20d1865a, points: 2, grade: true, solution: false -->
```python
# проверка правильности формирования холодных рекомендаций - 2 балл

assert isinstance(reco_cold, pd.DataFrame)

expected_columns = {"user_id", "item_id", "rank"}
assert expected_columns.issubset(set(reco_cold.columns))

assert reco_cold.user_id.nunique() == 111690

assert (
    reco_cold.groupby("user_id")["item_id"].nunique() == 10
).all(), "Ошибка: у каждого user_id должно быть ровно 10 уникальных item_id"
```

```python
MAP(k=10).calc(reco_cold, test)
```

Output:
```
0.07148914311412494
```

<!-- grade_id: cell-4c0fc9cecc555068, points: 2, grade: true, solution: false -->
```python
# проверка метрик качества холодных рекомендаций на полном test - 2 балла
assert MAP(k=10).calc(reco_cold, test) >= 0.04 and MAP(k=10).calc(reco_cold, test) < 0.98
```

# Рекомендации для горячих пользователей

Допишите метод `recommend` класса UseKnn так, чтобы он выдавал ровно k рекомендаций (сейчас он выдает некоторым пользователям меньше k, особенность implicit knn модели)

### Порекомендуйте k=10 уникальных айтемов для каждого горячего пользователя из тестовой выборки. Сохраните рекомендации в pandas DataFrame и назовите его `reco_hot`

датафрейм `reco_hot` ддолжен иметь обязательные поля `user_id`, `item_id`, `rank`. 

⚠️ горячими считаем пользователей, у которых ЕСТЬ любое количество интеракции в train. 

```python
train_users = set(train["user_id"].unique())
test_users = set(test["user_id"].unique())
hot_users = list(test_users.intersection(train_users))
print(len(hot_users))
```

Output:
```
154391
```

```python
reco_hot = user_knn.recommend(np.array(hot_users), k=10)
```

```python
reco_hot = reco_hot[[Columns.User, Columns.Item, Columns.Rank]].rename(
    columns={Columns.User: "user_id", Columns.Item: "item_id", Columns.Rank: "rank"}
)
```

```python
# reco_hot.to_csv(data_base_path / "reco_hot.csv", index=False)
```

## Hack в помощь

Дело к близится к вечеру, продакт менеджер торопит Вас.

А код userknn фитится долго, потому что кто-то запихнул в него весь train с over 900 тыс пользователей. 

Вы решаете аккуратно уменьшить трейн

- `возьмите только последние 30 дней датасета в трейн`


<!-- grade_id: cell-2df2fc8cf4a17ae8, points: 0, grade: false, solution: true -->
```python
max_train_date = train[Columns.Datetime].max()
new_train = train[train[Columns.Datetime] >= (max_train_date - pd.Timedelta(days=30))].copy()
new_train.shape
```

Output:
```
(1542040, 5)
```

```python
new_train_users = set(new_train["user_id"].unique())
test_users = set(test["user_id"].unique())
hot_users_new = list(test_users.intersection(new_train_users))
print(len(hot_users_new))
```

Output:
```
128070
```

Обучите модель, сделайте рекомендации `reco_hot` 

<!-- grade_id: cell-fb4fb28c7add6bc9, points: 0, grade: false, solution: true -->
```python
user_knn_new = UserKnn(model=ItemItemRecommender(), N_similar_users=50)
```

```python
user_knn_new.fit(new_train)
```

Output:
```
100%|██████████| 398800/398800 [00:41<00:00, 9598.37it/s]
```

```python
# import pickle

# models_base_path = Path("../models/").resolve()

# with open(models_base_path / "user_knn_model_small.pkl", "wb") as f:
#     pickle.dump(user_knn, f)

# with open(models_base_path / "user_knn_model_small.pkl", "rb") as f:
#     user_knn = pickle.load(f)
```

```python
reco_hot_new = user_knn_new.recommend(np.array(hot_users_new), k=10)
```

```python
reco_hot_new = reco_hot_new[[Columns.User, Columns.Item, Columns.Rank]].rename(
    columns={Columns.User: "user_id", Columns.Item: "item_id", Columns.Rank: "rank"}
)
```

```python
# reco_hot_new.to_csv(data_base_path / "reco_hot_new.csv", index=False)
```

```python
reco_hot = reco_hot_new.copy()
```

## Проверка рекомендаций по горячим пользователям - 8 баллов

Внимание! Есть скрытые тесты 

<!-- grade_id: cell-8c1fb777f3ad7af5, points: 3, grade: true, solution: false -->
```python
# проверки на правильность формирования reco_hot - 3 балла за все проверки в ячейке

assert isinstance(reco_hot, pd.DataFrame)

expected_columns = {"user_id", "item_id", "rank"}
assert expected_columns.issubset(set(reco_hot.columns))

assert reco_hot.user_id.nunique() == 128070
```

```python
print(MAP(k=10).calc(reco_hot, test))
print(MeanInvUserFreq(k=10).calc(reco_hot, test))
```

Output:
```
0.0019512742290850513
11.676641152462807
```

<!-- grade_id: cell-68ad1626e291a863, points: 5, grade: true, solution: false -->
```python
# проверки на метрики - 5 баллов за все проверки в ячейке

assert (
    reco_hot.groupby("user_id")["item_id"].nunique() == 10
).all(), "Ошибка: у каждого user_id должно быть ровно 10 уникальных item_id"

assert MAP(k=10).calc(reco_hot, test) <= 0.04 and MeanInvUserFreq(k=10).calc(reco_hot, test) > 7.0
```

# Как сдать ноутбук `knn.ipynb` на проверку

⚠️ Важное замечание: чтобы ваш ноутбук смог пройти проверку, скопируйте код класса `UserKnn` из `userknn.py` в этот ноутбук. Мы не можем гарантировать, что импорты из py файла будут работать.

1. Прогоните весь код ноутбука - проверьте, что нет ошибок и тесты проходят
2. Выложите готовый ноутбук в ваш репозиторий с сервисом из домашнего задания №1 по пути `notebooks/hw_3/knn.ipynb` в ветке `hw_3`

3. Проверьте, что есть доступ к вашему репозиторию для аккаунтов `https://github.com/feldlime`

5. Откройте PR в main ветку и добавьте в ревьюеры **своего ментора**

6. Не проводите мердж в `main` ветку, пока не увидите оценку за это ДЗ в ведомости. Файл с ноутбуком должен находиться в ветке `hw_3`

Обратите внимание, что сборка ноутбуков на проверку автоматизирована. В случае неправильного пути, имени файла или ветки (а также при отсутствии доступа у `@feldlime`) ваша работа не попадёт на проверку и получит `0` баллов.

# Баллы по ДЗ №3: максимум 25 баллов

1. прохождение проверки кода в ноутбуке `knn.ipynb` - **15 баллов**
2. обернуть модель UserKnn в сервис и побить безлайн `map@10 = 0.063` на лидерборде. Оценивается только лидерборд, без код ревью - **10 баллов**

## Комментарии  

- Вы можете переспользовать в своем сервисе код из `userknn.py`, который использован в ноутбуке `knn.ipynb`

Как реализовать модель в сервисе: 

- онлайн вариант: обучаете модель, сохраняете обученную модель (pickle, dill), при запуске сервиса ее поднимаете и запрашиваете рекомендации "на лету" 
- оффлайн вариант: предварительно посчитайте рекомендации для всех пользователей, сохраните и запрашивайте их
- в приватном тесте лидерборда есть как холодные, так и горячие пользователи